# Basics: Arithmetic, Assignment, and Strings

Topics:
1. Infix arithmetic operators
2. Numeric literals
3. Comparisons
4. Assignments
5. Strings

## Arithmetic should largely be familiar

```julia
3 + 7 # addition
```

```julia
10 - 3 # subtraction
```

```julia
20 * 5 # multiplication
```

```julia
100 / 10 # division
```

```julia
10 ^ 2 # exponentiation
```

```julia
101 % 2 # remainder (modulus)
```

```julia
sqrt(2) # square root
```

```julia
√2 # Unicode to the rescue: \sqrt + TAB
√(1+1)
```

Output:
```
1.4142135623730951
```

Note that dividing two integers yields a floating point number. There are two additional operators that may be helpful here:

```julia
10 / 6
```

Output:
```
1.6666666666666667
```

```julia
10 ÷ 6 # \div TAB or the `div` function
```

Output:
```
1
```

```julia
div(10, 6)
```

Output:
```
1
```

```julia
10 // 6
```

Output:
```
5//3
```

### Numbers: Many different ways to write the number forty-two

```julia
fortytwos = (42, 42.0, 4.20e1, 4.20f1, 84//2, 0x2a)
```

Output:
```
(42, 42.0, 42.0, 42.0f0, 42//1, 0x2a)
```

```julia
for x in fortytwos
    show(x)
    println("\tisa $(typeof(x))")
end
```

Output:
```
42	isa Int64
42.0	isa Float64
42.0	isa Float64
42.0f0	isa Float32
42//1	isa Rational{Int64}
0x2a	isa UInt8
```

### Bitwise arithmetic

```julia
0x2a & 0x70 # AND
```

Output:
```
0x20
```

```julia
0x2a | 0x70 # OR
```

Output:
```
0x7a
```

```julia
42 & 112
```

Output:
```
32
```

```julia
0b0010 << 2 # == 0b1000
```

Output:
```
0x08
```

### Logical operators

```julia
false && true # AND
```

Output:
```
false
```

```julia
false || true # OR
```

Output:
```
true
```

Note that they "short-circuit!"

```julia
x = -42
x > 0 || error("x must be positive")
```

Output:
```
@ Base .\error.jl:35

 [2] top-level scope

   @ c:\Users\jwei\Documents\GitHub\Julia\JuliaTutorials\introductory-tutorials\intro-to-julia\long-version\010 Arithmetic, Assignment, and Strings.ipynb:2
```

### Comparisons

```julia
1 == 1.0 # Equality
```

Output:
```
true
```

```julia
1 === 1.0 # Programmatically identical
```

Output:
```
false
```

```julia
3 < π
```

Output:
```
true
```

```julia
1 <= 1
```

```julia
.1 + .2
```

```julia
.1 + .2 ≈ .3 # \approx + TAB
```

Output:
```
true
```

Comparisons "chain"

Try inserting parentheses around one of these comparisons

```julia
2 == 2.0 == 0x02
```

Output:
```
true
```

```julia
(2 == 2.0) == 0x02
```

Output:
```
false
```

```julia
x = 42
0 < x < 100 || error("x must be between 0 and 100")
```

Output:
```
true
```

# Assignment

Assignment in Julia is done with the single `=`. All it does is associates a name (on the left) to a value (on the right).

```julia
x = 1 # Use the name `x` for the value `1`
```

Output:
```
1
```

```julia
y = x # Use the name `y` for the value that `x` refers to
```

Output:
```
1
```

```julia
x = "hello!" # Decide you have a better use for the name `x`
```

```julia
y # Is still the value 1
```

```julia
ϵ = eps(1.0) # You can make your own unicode names
```

Output:
```
2.220446049250313e-16
```

```julia
5ϵ # Juxtaposition is multiplication
```

Output:
```
1.1102230246251565e-15
```

## Updating operators

All the infix arithmetic operators above can be used as "updating" operators in conjunction with an assignment:

```julia
y += 1
```

Output:
```
2
```

```julia
# This is exactly the same as:
y = y + 1
```

Output:
```
3
```

Note that it's just re-purposing the _same name_ for a new value. This means that the type might even change:

```julia
y /= 2
```

Output:
```
1.5
```

# Strings

```julia
s1 = "I am a string."
```

```julia
s2 = """I am also a string. """
```

```julia
"Here, we get an "error" because it's ambiguous where this string ends "
```

```julia
"""Look, Mom, no "errors"!!! """
```

```julia
    println("""The other nice thing about triple-quoted
               string literals is that they ignore leading
               indentation, which is nice for long strings
               in real code. Try changing these quotes!""")
```

Output:
```
The other nice thing about triple-quoted
string literals is that they ignore leading
indentation, which is nice for long strings
in real code. Try changing these quotes!
```

Strings are not written with single `'`s — that's used for a single character:

```julia
first(s1)
```

```julia
'⊂'
```

```julia
'If you try writing a string in single-quotes, you will get an error'
```

## String interpolation

You can use the dollar sign inside a string to evaluate a Julia expression inside a string — either a single variable or a more complicated expression:

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

```julia
 println("That is $(num_fingers + num_toes) digits in all!!")
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```